Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

A Pythagorean triple is the set of positive integers that satisfy the equation $$a^2+b^2=c^2.$$ Given two numbers $a$ and $b$ with $a<b$ we can say these numbers are components of a Pythagorean triplet if $\sqrt{a^2+b^2}$ or $\sqrt{b^2-a^2}$ is an integer. <br/><br/>
The two problems you will be solving are the same problem given different ways.

Write a function LastTrip(a,b) that returns the last component of the Pythagorean triplet as integer if $a,b$ are components of a Pythagorean triplet else returns nothing.

<!-- grade_id: cell-63f48988a7135650, points: 0, grade: false, solution: true -->
```python
def LastTrip(a,b):
    if a > b:
        a,b = b,a
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2cdadfb7982b53fe, points: 5, grade: true, solution: false -->
```python
assert isinstance(LastTrip(3,4),int)
```

<!-- grade_id: cell-961b97b8f9f3f3f0, points: 5, grade: true, solution: false -->
```python
assert LastTrip(5,4) == 3
```

<!-- grade_id: cell-b6bb433eb55e89e1, points: 5, grade: true, solution: false -->
```python
assert LastTrip(39,80) == 89
```

<!-- grade_id: cell-08d8f549dda28620, points: 5, grade: true, solution: false -->
```python
assert LastTrip(10,15) is None 
```

Write a function isPTrip(a,b) that returns False if $a,b$ are not components of a Pythagorean triplet or returns (True,(a,b,c)) if $a,b$ are components of a Pythagorean triplet.
isPTrip(3,4) should return (True,(3,4,5))<br/>
isPTrip(3,5) should return (True,(3,4,5))<br/>
isPTrip(4,5) should return (True,(3,4,5))

<!-- grade_id: cell-a84c326e5204299f, points: 0, grade: false, solution: true -->
```python
def isPTrip(a,b):
    if a > b:
        a,b = b,a
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a1b0a5c3b3dab1f0, points: 3, grade: true, solution: false -->
```python
assert isinstance(isPTrip(3,4)[1],tuple)
```

<!-- grade_id: cell-586d3ba34e3d534f, points: 3, grade: true, solution: false -->
```python
assert isinstance(isPTrip(3,4)[0],bool)
```

<!-- grade_id: cell-680dd2bd8072c35b, points: 2, grade: true, solution: false -->
```python
assert isPTrip(3,4) == (True, (3,4,5))
```

<!-- grade_id: cell-b27ddf551a550228, points: 3, grade: true, solution: false -->
```python
assert isPTrip(8,15)==(True, (8, 15, 17))
```

<!-- grade_id: cell-64b8afeb20564ef5, points: 3, grade: true, solution: false -->
```python
assert isPTrip(80,89) == (True, (39, 80, 89))
assert isPTrip(3,5)==(True, (3,4,5))
```

<!-- grade_id: cell-0808904277a9adb0, points: 3, grade: true, solution: false -->
```python
assert isinstance(isPTrip(24,55),bool)
```

<!-- grade_id: cell-ff56b72812531039, points: 3, grade: true, solution: false -->
```python
assert isPTrip(24,55) == False
```