Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "MR. GERSHON TETTEH YAW AMANOR"
COLLABORATORS = "THE YAWS"
```

---

As learnt in level 100, the nth root of unity is defined by the expression 
$$z_i = \cos(\frac{2k\pi}{n})+i\sin(\frac{2k\pi}{n}), \ \mathsf{with}\ k = 0,1,2,\ldots,n-1$$
Write a function runity(n) that returns the nth root of unity. Your function shoult return a list of the roots of unity.

<!-- grade_id: cell-f25a7b8b5defce81, points: 0, grade: false, solution: true -->
```python
def runity(n):
    # YOUR CODE HERE
    from math import cos,sin,pi
    res = []
    k = range(0,n)
    for i in range(1,n+1):
        z = cos((2*k[i-1]*pi)/n) + 1j*sin((2*k[i-1]*pi)/n)
        res.append(z)
        i += 1
    return(res)    
    raise NotImplementedError()
```

<!-- grade_id: cell-c8410abc4b7c5e14, points: 2, grade: true, solution: false -->
```python
assert isinstance(runity(12),list)
```

<!-- grade_id: cell-f068f93014146124, points: 2, grade: true, solution: false -->
```python
assert isinstance(runity(1),list)
```

<!-- grade_id: cell-a89cdd6b1ea8010e, points: 2, grade: true, solution: false -->
```python
assert isinstance(runity(5)[0],complex)
```

<!-- grade_id: cell-35bc30bd34ce0df2, points: 2, grade: true, solution: false -->
```python
assert len(runity(12)) == 12
```

<!-- grade_id: cell-133277ba08c3ea44, points: 2, grade: true, solution: false -->
```python
assert runity(12)[0] == 1+1j*0
```

<!-- grade_id: cell-c9a70cdfa75520e6, points: 5, grade: true, solution: false -->
```python

```