Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "GERSHON TETTEH AMANOR"
COLLABORATORS = ""
```

---

### Please read all instructions carefully for each cell. Variable names must be defined as strictly as the question dictates. 
### Enter your solution in the space marked "YOUR CODE HERE"
Failure to do this will result in loss of marks for that question and maybe even subsequent questions.

```python
import numpy as np
```

Create a discrete interval $x \in \left[ -2,2 \right]$ with 201 discrete points in the interval

<!-- grade_id: cell-8f155b3499db3426, points: 0, grade: false, solution: true -->
```python
### Create your discrete interval x here
# YOUR CODE HERE
x = np.linspace(-2,2,201)
```

<!-- grade_id: cell-67e6e02fe1996264, points: 5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert isinstance(x,np.ndarray)
assert x.shape==(201,)
### END TESTS
```

<!-- grade_id: cell-5a094ca6b16261d9, points: 5, grade: true, solution: false -->
```python

```

Using the $x$ created above, create the Gaussian $y$ given by $$y = \frac{1}{\sigma\sqrt{2\pi}}e^{\displaystyle{-\frac{(x-\mu)^2}{2\sigma^2}}}$$. Use $\mu = 2$ and $\sigma = 1$

<!-- grade_id: cell-576338ffe2deba80, points: 0, grade: false, solution: true -->
```python
### Note that the y created here is an array
# YOUR CODE HERE
𝜇=2
𝜎=1
y = (1/(𝜎*np.sqrt(2*np.pi)))*(np.exp(-(x-𝜇)**2/((2*𝜎)**2)))
```

<!-- grade_id: cell-aab71fe9cc1ab85f, points: 5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert isinstance(y,np.ndarray)
assert y.shape==(201,)
### END TESTS
```

<!-- grade_id: cell-acd11bb0210e5db1, points: 5, grade: true, solution: false -->
```python

```

Create a discrete interval **x2** with $x2 \in \left[ -4,1 \right]$ with 201 discrete points in the interval

<!-- grade_id: cell-ea02381e1f341346, points: 0, grade: false, solution: true -->
```python
### Create your discrete interval x2 here
# YOUR CODE HERE
x2 = np.linspace(-4,1,201)
```

<!-- grade_id: cell-ec86bf86c343bd79, points: 5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert x2.min() == -4
assert x2.max() == 1
### END TESTS
```

<!-- grade_id: cell-12758a1cc5fd9088, points: 5, grade: true, solution: false -->
```python

```

In the next question, we are going to create a step function. We define the array **y_3** as 
$$y_3(x) = \lfloor x^2 + 3x + 5 \rfloor$$ where $\lfloor \rfloor$ is the floor function.

<!-- grade_id: cell-eb5e0a75d60da9b9, points: 0, grade: false, solution: true -->
```python
### your y_3 should be an array
# YOUR CODE HERE
from math import floor
y_3 = np.array([floor((x**2 + 3*x +5)) for x in x2])
```

<!-- grade_id: cell-3f31672c58730c0a, points: 5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert isinstance(y_3,np.ndarray)
assert y_3.shape==(201,)
### END TESTS
```

<!-- grade_id: cell-b7582f0e1362ae6f, points: 5, grade: true, solution: false -->
# $\left[*hey*\right]$

$$EJ(x,y)\ =\ \frac{X.Y}{\||X||.||Y||-\ X.Y}$$

```python
def isp(x):
    impoort numpy as np
    if x== 0 or 1:
        print(x,' is not prime')
    elif x == 2:
        print(x, ' is prime')
    elif x>2:
        pc = []
        i = list(range(2,x))
        p = list(np.ones(x-2, dtype=int)*x)
        for j in range(x-2):
            pc.append(p[j]%i[j])
            print(pc)
        if (0 in pc)==False:
            print(x,' is prime')
        else:
            print(x,' is not prime')
```

Output:
```
Running cells with 'c:\Users\user\AppData\Local\Programs\Python\Python311\python.exe' requires the ipykernel package.

Run the following command to install 'ipykernel' into the Python environment. 

Command: 'c:/Users/user/AppData/Local/Programs/Python/Python311/python.exe -m pip install ipykernel -U --user --force-reinstall'
```

```python

```