Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "GERSHON TETTEH AMANOR"
COLLABORATORS = ""
```

---

```python
from numpy import isclose
```

One of the integrals that is quite complicated is the error function $$\frac{2}{\sqrt{\pi}}\int_0^x e^{-t^2}dt=\mathrm{erf}(x).$$ It is however commonly used in probability, statistics and the theory of partial differential equations. We are going to approximate the error function using the Riemann sum. Using Riemann sums, we can approximate the definite integral $$\int_a^bf(x)\approx \sum_{i=1}^n f(x_i^*)\Delta x_i$$ where the interval $\left[ a,b \right]$ is partitioned into the points $a=x_0<x_1<x_2\ldots<x_n=b$ and $\Delta x_i=x_i-x_{i-1}$ and $x_i^*\in\left[ x_{i-1},x_i \right].$ This essentially approximates the integral using the sum of rectangles $$\sum R_i$$ where $R_i$ has height $f(x_i^*)$ and width $\Delta x_i.$ 

Write a function erf(x,n) that approximates the error function at the point $x$ using $n$ rectangles ($n+1$ points). We will choose $x_i^*$ to be the left endpoint $x_{i-1}.$ Remember to scale your result by $\frac{2}{\sqrt{\pi}}.$ Note: you can use linspace from numpy to perform the partitioning of the interval.

<!-- grade_id: cell-098ec09b54417eed, points: 0, grade: false, solution: true -->
```python
import numpy as np

def erf(x,n):
    # YOUR CODE HERE
    deltax = x/n
    xi = np.linspace(0,x,n+1)
    fxi = np.exp(-xi**2)
    return 2/np.sqrt(np.pi)*np.sum(fxi[:-1]*deltax)
    raise NotImplementedError()
```

```python
erf(2,5)
```

Output:
```
1.2157918631319096
```

```python
erf(2,5)
```

Output:
```
1.2157918631319096
```

# Alternative

```python
def erf(x,n):
    from numpy import linspace
    from math import exp, pi,sqrt
    x = linspace(0,x,n+1)
    res = 0
    for i in range(x.size-1):
        res += exp(-x[i]**2)*(x[i+1]-x[i])
    return res*2/sqrt(pi)
```

x,n=2,5
xi = np.linspace(0,x,n+1)
xi
deltax = x/n
deltax
delta = [(xi[i]-xi[i-1]) for i in range(len(xi))]
fxi = np.exp(-xi**2)
fxi
2/np.sqrt(np.pi)*np.sum(fxi[:]*delta)

```python
erf(2,5)
```

Output:
```
1.2157918631319096
```

```python
erf(1,1001)
```

Output:
```
0.8430570034585734
```

```python
0.8430570034585734
```

Output:
```
0.8430570034585734
```

$\rightarrow$

$\leftarrow$

$\downarrow$

$\frac{2}{\sqrt{\pi}}\int_0^x e^{-t^2}dt=\mathrm{erf}(x).$

```python

```

<!-- grade_id: cell-fac4ea36575424e0, points: 2, grade: true, solution: false -->
```python
assert isclose(erf(1,1001),0.8430570034585734)
```

<!-- grade_id: cell-5c060e3768578b0d, points: 2, grade: true, solution: false -->
```python
assert isclose(erf(-5,101),-1.0279301774013976)
```

<!-- grade_id: cell-1ecd238c5d802aef, points: 2, grade: true, solution: false -->
```python
assert isclose(erf(100,501),1.1126126913268979)
```

A sequence $\left\lbrace x_i \right\rbrace_{i=1}^n$ is said to be decreasing if $x_{i+1}<x_i$ for all $i=0,1,\ldots,n.$ Write a function isDecreasing(x) that takes as input a list x of numbers and returns True if the list represents a decreasing sequence otherwise returns False.

<!-- grade_id: cell-b5294b3c2316cc8c, points: 0, grade: false, solution: true -->
```python
def isDecreasing(x):
    # YOUR CODE HERE
    for i in range(len(x) - 1):
        if x[i] <= x[i + 1]:
            return False
    return True
    raise NotImplementedError()
```

# $$My code$$

```python
def isDecreasing(x):
    # YOUR CODE HERE
    for i in range(0,len(x)-1):
        if x[i+1]<x[i]:
            dec = True
    return dec         
    raise NotImplementedError()
```

# $$Alternative$$

```python
def isDecreasing(x):
    res = [x[i]<x[i-1] for i in range(1,len(x))]
    return all(res)
```

<!-- grade_id: cell-63e1fc1828c905f9, points: 2, grade: true, solution: false -->
```python
assert isDecreasing([9,3,0,-2,-9,-15]) == True
```

<!-- grade_id: cell-5fe3fc025abde364, points: 1, grade: true, solution: false -->
```python
assert isDecreasing([9,3,0,-2,-2,-15]) == False
```

<!-- grade_id: cell-45eda5efa377407b, points: 1, grade: true, solution: false -->
```python
assert isDecreasing([9,3,0,-2,3,-15]) == False
```

Given the function $f(x),$ we can approximate the arc length from $(a,f(a))$ to $(b,f(b))$ as the sum $$\sum_{i=1}^n\sqrt{{\Delta x_i}^2+{\Delta y_i}^2}$$ where the interval $\left[ a,b \right]$ is partitioned into the points $a=x_0<x_1<x_2\ldots<x_n=b$ and $\Delta x_i=x_i-x_{i-1}$ and $\Delta y_i=f(x_i)-f(x_{i-1}).$ Write a function arclength(f,a,b,n) that takes f a function, the interval $[a,b]$, the number of subintervals n (or n+1 points) and returns an approximation of the arc length. You can use linspace from numpy to create your partition.

# $$My code$$

```python

def arclength(f,a,b,n):
     # YOUR CODE HERE
        import numpy as np
        x = np.linspace(a,b,n+1)
        y = f(x)
        dx = x[1]-x[0]
        dy = y[1]-y[0]
        AL = np.sqrt(dx**2 + dy**2)
        for i in range(1,len(x)-1):
            dx = x[i+1]-x[i]
            dy = y[i+1]-y[i]
            AL += np.sqrt(dx**2 + dy**2)
        return AL      
        raise NotImplementedError()
```

```python

```

<!-- grade_id: cell-4c3078b4e9357130, points: 0, grade: false, solution: true -->
```python

def arclength(f,a,b,n):
     # YOUR CODE HERE
    x = np.linspace(a,b,n+1)
    y = f(x)
    dx = x[1]-x[0]
    dy = y[1]-y[0]
    L = np.sqrt(dx**2 + dy**2)
    for i in range(1,n):
        dx = x[i+1]-x[i]
        dy = y[i+1]-y[i]
        L += np.sqrt(dx**2 + dy**2)
    return L
    raise NotImplementedError()
```

# $$Alternative$$

```python
def arclength(f,a,b,n):
    from numpy import linspace
    from math import sqrt
    x = linspace(a,b,n+1)
    dx = (x[1]-x[0])**2
    y = f(x)
    res = 0
    for i in range(1,len(x)):
        res += sqrt(dx+(y[i]-y[i-1])**2)
    return res    
```

<!-- grade_id: cell-f4da8797c6a798e9, points: 3, grade: true, solution: false -->
```python
assert isclose(arclength(lambda x: x**2,1,3,100001),8.268145901058078)
```

<!-- grade_id: cell-7ddc612d2ecbc020, points: 4, grade: true, solution: false -->
```python
assert isclose(arclength(lambda x: x**0.5,1,3,1001),2.1325089836938407)
```

<!-- grade_id: cell-680c58f7228a9438, points: 3, grade: true, solution: false -->
```python
assert isclose(arclength(lambda x: 5*x-2,-2,3,1001),25.49509756796383)
```

# $$[AMA]_{left}^{right}$$

```python

```