Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "GERSHON TETTEH AMANOR"
COLLABORATORS = ""
```

---

Python supports the concept of nested lists. One potential use of this concept is as matrices. For instance the following code
```
A = [[a,b],[c,d]]
```
can represent the matrix $$A=\begin{bmatrix} a & b \\ c & d \end{bmatrix}.$$ It is therefore easy to see that 
```
A[0][0]=a,
A[0][1]=b,
A[1][0]=c,
A[1][1]=d
```

Given the matrix $$A=\begin{bmatrix} a & b \\ c & d \end{bmatrix}$$ we know that the determinant of $A$ is given as $$\det(A)=ad-bc.$$ Write a function det(A) that takes as input A a 2x2 matrix written as a list (nested) and returns the determinant of the matrix represented by the list A. Your code should ensure that your matrix is 2x2 and return none if that condition is violated

<!-- grade_id: cell-63f48988a7135650, points: 0, grade: false, solution: true -->
```python
def det(A):
    # YOUR CODE HERE
    # we check if each of the nested lists and the external list are of length 2
    if len(A)==2:
        D =(A[0][0]*A[1][1])-(A[0][1]*A[1][0])
    else:
        D = None
    return(D)
    raise NotImplementedError()
```

<!-- grade_id: cell-2cdadfb7982b53fe, points: 2, grade: true, solution: false -->
```python
assert isinstance(det([[1,0],[0,1]]),int)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[1], line 1
----> 1 assert isinstance(det([[1,0],[0,1]]),int)

NameError: name 'det' is not defined
```

<!-- grade_id: cell-961b97b8f9f3f3f0, points: 2, grade: true, solution: false -->
```python
assert det([[0,1],[1,0]]) == -1
```

<!-- grade_id: cell-b6bb433eb55e89e1, points: 2, grade: true, solution: false -->
```python
assert det([[5,0],[2,4]]) == 20
```

<!-- grade_id: cell-08d8f549dda28620, points: 2, grade: true, solution: false -->
```python
A = [[4,5],[-8,5]]
assert det(A) == A[0][0]*A[-1][-1] - A[0][-1]*A[-1][0]
```

The transpose of the matrix $$A=\begin{bmatrix} a & b \\ c & d \end{bmatrix}$$ is the matrix $$A=\begin{bmatrix} a & c \\ b & d \end{bmatrix}.$$ Write a function T(A) that takes as input a matrix A written as a nested list and returns the transpose as a nested list. 

<!-- grade_id: cell-6e899888de3c5c07, points: 0, grade: false, solution: true -->
```python
def T(A):
    # YOUR CODE HERE
    if len(A)==2:
        T =[[A[0][0],A[1][0]],[A[0][1],A[1][1]]]
    else:
        T = None
    return(T)
    raise NotImplementedError()
```

<!-- grade_id: cell-b6f9b892e0cf798a, points: 2, grade: true, solution: false -->
```python
assert T([[1,5],[2,3]]) == [[1, 2], [5, 3]]
```

<!-- grade_id: cell-09cc9f6b625f6aa3, points: 3, grade: true, solution: false -->
```python
assert T([[7,-1],[5,4]]) == [[7, 5], [-1, 4]]
```

We can find the inverse of a 2x2 matrix $$A=\begin{bmatrix} a & b \\ c & d \end{bmatrix}$$ as $$A^{-1}=\begin{bmatrix} \frac{d}{det(A)} & \frac{-b}{det(A)} \\ \frac{-c}{det(A)} & \frac{a}{det(A)} \end{bmatrix}.$$ Write a function inv(A) that takes as input a 2x2 matrix A written as a list (nested) and returns its inverse as a nested list.

<!-- grade_id: cell-a84c326e5204299f, points: 0, grade: false, solution: true -->
```python
def inv(A):
    # YOUR CODE HERE
    dt = det(A)
    #if len(A)==2:
    I=[[A[1][1]/dt,-1*A[0][1]/dt],[-1*A[1][0]/dt,A[0][0]/dt]]
    #else:
    #   I = None
    return(I)    
    raise NotImplementedError()
```

<!-- grade_id: cell-a1b0a5c3b3dab1f0, points: 3, grade: true, solution: false -->
```python
assert abs(det(inv([[4,5],[-8,5]]))*det([[4,5],[-8,5]])-1)<1e9
```

<!-- grade_id: cell-586d3ba34e3d534f, points: 2, grade: true, solution: false -->
```python
assert inv([[1,0],[0,1]]) == [[1,0],[0,1]]
```

<!-- grade_id: cell-680dd2bd8072c35b, points: 2, grade: true, solution: false -->
```python
assert inv([[0,1],[1,0]]) == [[0.0, 1.0], [1.0, 0.0]]
```

Given the 2x2 matrix $$A=\begin{bmatrix} a_{11} & a_{12} \\ a_{21} & a_{22} \end{bmatrix}$$ and the vector $\boldsymbol{b}$, $$B=\begin{bmatrix} b_{1}\\ b_{2} \end{bmatrix},$$ the product $A\boldsymbol{b}$ is given as
$$A\boldsymbol{b}=\begin{bmatrix} a_{11}b_{1}+a_{12}b_{2} \\ a_{21}b_{1}+a_{22}b_{2} \end{bmatrix}.$$ 
Write a function mvprod(A,b) that takes as input a 2x2 matrices A written as nested list, and a two element vector b written as a simple list and returns their product as a simple list.

<!-- grade_id: cell-83164630e3e9df22, points: 0, grade: false, solution: true -->
```python
def mvprod(A,b):
    # YOUR CODE HERE
    prod = [A[0][0]*b[0] + A[0][1]*b[1], A[1][0]*b[0] + A[1][1]*b[1]]
    return(prod)
    raise NotImplementedError()
```

<!-- grade_id: cell-9662f03b885450da, points: 3, grade: true, solution: false -->
```python
assert mvprod([[1,2],[1,4]],[2,3]) == [8, 14]
```

<!-- grade_id: cell-71fd9a50b3d7c4b7, points: 2, grade: true, solution: false -->
```python
assert mvprod([[2,1],[3,0]],[4,1]) == [9, 12]
```

Given the 2x2 matrices $$A=\begin{bmatrix} a_{11} & a_{12} \\ a_{21} & a_{22} \end{bmatrix}$$ and $$B=\begin{bmatrix} b_{11} & b_{12} \\ b_{21} & b_{22} \end{bmatrix},$$ the product $AB$ is given as
$$AB=\begin{bmatrix} a_{11}b_{11}+a_{12}b_{21} & a_{11}b_{12}+a_{12}b_{22} \\ a_{21}b_{11}+a_{22}b_{21} & a_{21}b_{12}+a_{22}b_{22} \end{bmatrix}.$$ 
Note that $AB$ can also written as $$AB = \begin{bmatrix} A\boldsymbol{b}_1 & A\boldsymbol{b}_2 \end{bmatrix}$$ where $$\boldsymbol{b}_k = \begin{bmatrix} b_{1k} & \ldots & b_{nk} \end{bmatrix}^T$$
Write a function mprod(A,B) that takes as input two 2x2 matrices A, B written as nested lists and returns their product as a nested list.

<!-- grade_id: cell-4eb318b61babe2dc, points: 0, grade: false, solution: true -->
```python
def mprod(A,B):
    # YOUR CODE HERE
    prod = [[A[0][0]*B[0][0] + A[0][1]*B[1][0], A[0][0]*B[0][1] + A[0][1]*B[1][1]], [A[1][0]*B[0][0] + A[1][1]*B[1][0], A[1][0]*B[0][1] + A[1][1]*B[1][1]]]
    return(prod)
    raise NotImplementedError()
```

# my alternative try
def mprod(A,B):
    # YOUR CODE HERE
    prod =[mvprod(A,A[0]),mvprod(A,A[1])]
    return(prod)
    raise NotImplementedError()

<!-- grade_id: cell-5672faf68ac2f9e1, points: 5, grade: true, solution: false -->
```python
assert mprod([[1,2],[1,4]],[[2,1],[3,0]]) == [[8, 1], [14, 1]]
```

<!-- grade_id: cell-c10a193a2c4862fa, points: 5, grade: true, solution: false -->
```python
assert mprod([[1,2],[1,4]],inv([[1,2],[1,4]])) == [[1.0, 0.0], [0.0, 1.0]]
```

# To God be the Glory For ever and ever

```python
A =[[1,2],[3,4]]
```

```python
det(A)
```

```python
T(A)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[24], line 1
----> 1 T(A)

NameError: name 'T' is not defined
```

```python
inv(A)
```

```python
mvprod(A,[1,2])
```

```python
mprod(A,A)
```

Output:
```
[[-24, 45], [-72, -15]]
```

```python
from cmath import e
a = e**2j
a=complex(a)
print(a)
a.imag
```

Output:
```
(-0.4161468365471424+0.9092974268256817j)
0.9092974268256817
```

```python

```